---
title: 'R Data Visualization Packages in R'
description: 'Introduce various data visualization packages in R.'
free_preview: true
---

## Example coding exercise

```yaml
type: NormalExercise
key: 2bafef99a3
lang: r
xp: 100
skills: 1
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---

## Data Visualization Packages in R

```yaml
type: VideoExercise
key: ec1308af30
xp: 50
```

`@projector_key`
a04e0c8726256b0922bee5a00df24548
